Partitioning changes in ecosystem productivity by effects of species interactions in biodiversity experiments

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public review):

As a starting point, the authors discuss the so-called "additive partitioning" (AP) method proposed by Loreau & Hector in 2001. The AP is the result of a mathematical rearrangement of the definition of overyielding, written in terms of relative yields (RY) of species in mixtures relative to monocultures. One term, the so-called complementarity effect (CE), is proportional to the average RY deviations from the null expectations that plants of both species "do the same" in monocultures and mixtures. The other term, the selection effect (SE), captures how these RY deviations are related to monoculture productivity. Overall, CE measures whether relative biomass gains differ from zero when averaged across all community members, and SE, whether the "relative advantage" species have in the mixture, is related to their productivity. In extreme cases, when all species benefit, CE becomes positive.

This is not true; positive CE does not require positive RY deviations of all species. CE is positive as long as average RY deviation is greater than 0. In a 2-species mixture, for example, if the RY deviation of one species is -0.2 and that of the other species is +0.3, CE would be still positive. Positive CE can be associated with negative NE (net biodiversity effects) when more productivity species have smaller negative RY deviation compared to positive RY deviation of less productive species. Therefore, the suggestion by the reviewer “This is intuitively compatible with the idea that niche complementarity mitigates competition (CE>0)” is not correct.

When large species have large relative productivity increases, SE becomes positive. This is intuitively compatible with the idea that niche complementarity mitigates competition (CE>0), or that competitively superior species dominate mixtures and thereby driver overyielding (SE>0).

The use of word “mitigate” indicates that the effects of niche complementarity and competition are in opposite directions, which is not true with biodiversity experiments based on replacement design. We have explained this in detail in our first responses to reviewers.

However, it is very important to understand that CE and SE capture the "statistical structure" of RY that underlies overyielding. Specifically, CE and SE are not the ultimate biological mechanisms that drive overyielding, and never were meant to be. CE also does not describe niche complementarity. Interpreting CE and SE as directly quantifying niche complementarity or resource competition, is simply wrong, although it sometimes is done. The criticism of the AP method thus in large part seems unwarranted. The alternative methods the authors discuss (lines 108-123) are based on very similar principles.

Agree. However, If CE and SE are not meant to be biological mechanisms, as suggested by the reviewer, the argument “This is intuitively compatible with the idea that niche complementarity mitigates competition (CE>0), or that competitively superior species dominate mixtures and thereby driver overyielding (SE>0)” would be invalid.

Lines 108-123 are not on our method.

The authors now set out to develop a method that aims at linking response patterns to "more true" biological mechanisms.

Assuming that "competitive dominance" is key to understanding mixture productivity, because "competitive interactions are the predominant type of interspecific relationships in plants", the authors introduce "partial density" monocultures, i.e. monocultures that have the same planting density for a species as in a mixture. The idea is that using these partial density monocultures as a reference would allow for isolating the effect of competition by the surrounding "species matrix".

The authors argue that "To separate effects of competitive interactions from those of other species interactions, we would need the hypothesis that constituent species share an identical niche but differ in growth and competitive ability (i.e., absence of positive/negative interactions)." - I think the term interaction is not correctly used here, because clearly competition is an interaction, but the point made here is that this would be a zero-sum game.

We did not say that competition is not an interaction.

The authors use the ratio of productivity of partial density and full-density monocultures, divided by planting density, as a measure of "competitive growth response" (abbreviated as MG). This is the extra growth a plant individual produces when intraspecific competition is reduced.

Here, I see two issues: first, this rests on the assumption that there is only "one mode" of competition if two species use the same resources, which may not be true, because intraspecific and interspecific competition may differ. Of course, one can argue that then somehow "niches" are different, but such a niche definition would be very broad and go beyond the "resource set" perspective the authors adopt. Second, this value will heavily depend on timing and the relationship between maximum initial growth rates and competitive abilities at high stand densities.

True. Research findings indicate that biodiversity effect detected with AP is not constant.

The authors then progress to define relative competitive ability (RC), and this time simply uses monoculture biomass as a measure of competitive ability. To express this biomass in a standardized way, they express it as different from the mean of the other species and then divide by the maximum monoculture biomass of all species.

I have two concerns here: first, if competitive ability is the capability of a species to preempt resources from a pool also accessed by another species, as the authors argued before, then this seems wrong because one would expect that a species can simply be more productive because it has a broader niche space that it exploits. This contradicts the very narrow perspective on competitive ability the authors have adopted. This also is difficult to reconcile with the idea that specialist species with a narrow niche would outcompete generalist species with a broad niche.

Competitive ability is not necessarily associated with species niche space. Both generalist and specialist species can be more productive at a particular study site, as long as they are more capable of obtaining resources from a local pool. Remember, biodiversity experiments are conducted at a site of particular conditions, not across a range of species niche space at landscape level.

Second, I am concerned by the mathematical form. Standardizing by the maximum makes the scaling dependent on a single value.

As explained in lines 370-376, the mathematical form is a linear approximation as the relationship between competitive growth responses and species relative competitive ability is generally unknow but would be likely nonlinear. Once the relationship is determined in future research, the scaling factor is not needed.

As a final step, the authors calculate a "competitive expectation" for a species' biomass in the mixture, by scaling deviations from the expected yield by the product MG ⨯ RC. This would mean a species does better in a mixture when (1) it benefits most from a conspecific density reduction, and (2) has a relatively high biomass.

Put simply, the assumption would be that if a species is productive in monoculture (high RC), it effectively does not "see" the competitors and then grows like it would be the sole species in the community, i.e. like in the partial density monoculture.

Overall, I am not very convinced by the proposed method.

Comments on revised version:

Only minimal changes were made to the manuscript, and they do not address the main points that were raised.

Reviewer #2 (Public review):

This manuscript by Tao et al. reports on an effort to better specify the underlying interactions driving the effects of biodiversity on productivity in biodiversity experiments. The authors are especially concerned with the potential for competitive interactions to drive positive biodiversity-ecosystem functioning relationships by driving down the biomass of subdominant species. The authors suggest a new partitioning schema that utilizes a suite of partial density treatments to capture so-called competitive ability. While I agree with the authors that understanding the underlying drivers of biodiversity-ecosystem functioning relationships is valuable - I am unsure of the added value of this specific approach for several reasons.

No responses.

Comments on revised version:

The authors changed only one minor detail in response to the last round of reviews.

Reviewer #3 (Public review):

Summary:

This manuscript claims to provide a new null hypothesis for testing the effects of biodiversity on ecosystem functioning. It reports that the strength of biodiversity effects changes when this different null hypothesis is used. This main result is rather inevitable. That is, one expects a different answer when using a different approach. The question then becomes whether the manuscript's null hypothesis is both new and an improvement on the null hypothesis that has been in use in recent decades.

Our approach adopts two hypotheses, null hypothesis that is also with the additive partitioning model and competitive hypothesis that is new. Null hypothesis assumes that inter- and intra-specie interactions are the same, while competitive hypothesis assumes that species differ in competitive ability and growth rate. Therefore, our approach is an extension of current approach. Our approach separates effects of competitive interactions from those of other species interactions, while the current approach does not.

Strengths:

In general, I appreciate studies like this that question whether we have been doing it all wrong and I encourage consideration of new approaches.

Weaknesses:

Despite many sweeping critiques of previous studies and bold claims of novelty made throughout the manuscript, I was unable to find new insights. The manuscript fails to place the study in the context of the long history of literature on competition and biodiversity and ecosystem functioning.

We have explained in our first responses that competition and biodiversity effects are studied in different experimental approaches, i.e., additive and replacement designs. Results from one approach are not compatible with those from the other. For example, competition effect with additive design is negative but generally positive with replacement design that is used extensively in biodiversity experiments. We have considered species competitive ability, density-growth relationship, and different effects of competitive interactions between additive and replacement design, while the current method does not reflect any of those.

The Introduction claims the new approach will address deficiencies of previous approaches, but after reading further I see no evidence that it addresses the limitations of previous approaches noted in the Introduction. Furthermore, the manuscript does not reproducibly describe the methods used to produce the results (e.g., in Table 1) and relies on simulations, claiming experimental data are not available when many experiments have already tested these ideas and not found support for them.

We used simulation data, as partial density monocultures are generally not available in previous biodiversity experiments.

Finally, it is unclear to me whether rejecting the 'new' null hypothesis presented in the manuscript would be of interest to ecologists, agronomists, conservationists, or others.

Our null hypothesis is the same as the null hypothesis with the additive partitioning assuming that inter- and intra-species interactions are the same, while our competitive hypothesis assumes that species differ in competitive ability and growth rate. Rejecting null hypothesis means that inter- and intra-species interactions are different, whereas rejecting competitive hypothesis indicates existence of positive/negative species interactions. This would be interesting to everyone.

Comments on revised version:

Please see review comments on the previous version of this manuscript. The authors have not revised their manuscript to address most of the issues previously raised by reviewers.

No responses.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Do take reviews seriously. Even if you think the reviewers all are wrong and did not understand your work, then this seems to indicate that it was not clearly presented.

Reviewer #2 (Recommendations for the authors):

I can understand that the authors are perhaps frustrated with what they perceive as a basic misunderstanding of their goals and approach. This misunderstanding however, provides with it an opportunity to clarify. I believe that the authors have tried to clarify in rebutting our statements but would do better to clarify in the manuscript itself. If we reviewers, who are deeply invested in this field, don't understand the approach and its value, then it is likely that many readers will not as well.

The additive partitioning has been publicly questioned at least for serval times since the conception of the method in 2001. Our work provides an alternative.

Changes from the earlier version to the revised version

Abstract

Abstract

Species interactions affect ecosystem productivity. Positive interactions (resource partitioning and facilitation) increase productivity while negative interactions (species interference) decrease productivity relative to the null expectations defined by monoculture yields. Effects of competitive interactions (resource competition) can be either positive or negative. Distinguishing effects of species interactions is therefore difficult, if not impossible, with current biodiversity experiments involving mixtures and full density monocultures.

To partition changes in ecosystem productivity by effects of species interactions, we modify null expectations with competitive growth responses, i.e., proportional changes in individual size (biomass or volume) expected in mixture based on species differences in growth and competitive ability. We use partial density (species density in mixture) monocultures and the competitive exclusion principle to determine maximum competitive growth responses and full density monoculture yields to measure species ability to achieve maximum competitive growth responses in mixture. Deviations of observed yields from competitive expectations represent the effects of positive/negative species interactions, while the differences between competitive and null expectations reflect the effects of competitive interactions.

We demonstrate the effectiveness of our competitive partitioning model in distinguishing effects of species interactions using both simulated and experimental species mixtures. Our competitive partitioning model enables meaningful assessments of species interactions at both species and community levels and helps disentangle underlying mechanisms of species interactions responsible for changes in ecosystem productivity and identify species mixtures that maximize positive effects.

Species interactions and ecosystem productivity

Detecting changes in ecosystem productivity with biodiversity and understanding how these changes are affected by mechanisms of species interactions has been a central focus in ecology for decades (Loreau, 2000; Loreau and Hector, 2019; Mahaut et al., 2020). However, effective methods are lacking (Barry et al., 2019; Bourrat et al., 2023; Drake, 2003; Forrester and Pretzsch, 2015; Jaillard et al., 2018; Loreau and Hector, 2019). The method of additive partitioning was developed from the covariance equation to quantify changes in ecosystem productivity by comparing species performance in mixture with those expected from their monocultures using a replacement design (Barry et al., 2019; Loreau and Hector, 2001). The resultant changes in ecosystem productivity (commonly referred to as net biodiversity effects) are mathematically decomposed into additive components, namely the complementarity effect (CE) measuring community average response and the selection effect (SE) measuring interspecies differences (Forrester and Pretzsch, 2015; Loreau and Hector, 2001, 2019). CE ad SE are widely used to examine mechanisms of species interactions responsible for biodiversity effects, i.e., positive CE for positive interactions and positive SE for competitive interactions (Barry et al., 2019; Feng et al., 2022; Hagan et al., 2021; Loreau and Hector, 2001; Montès et al., 2008; Petchey, 2003; Polley et al., 2003), despite that neither of the additive components can be linked to specific effects of species interactions (Bourrat et al., 2023; Forrester and Pretzsch, 2015; Hagan et al., 2021; Hooper et al., 2005; Loreau et al., 2012; Loreau and Hector, 2019; Montès et al., 2008; Petchey, 2003).

Among the interspecific interactions in plants, positive interactions (resource partitioning and facilitation or collectively referred to as complementarity) increase productivity and negative interactions (species interference or interference competition) decrease productivity relative to the null expectations, i.e., species yields expected in mixture based on composition (relative yield) and full density monoculture yields (Barry et al., 2019; Hooper et al., 2005; Loreau, 2000; Loreau and Hector, 2001). Uncertainty comes with effects of competitive interactions (i.e., resource competition) that are positive for more competitive species and negative for less competitive species (Aschehoug et al., 2016; Barry et al., 2019). For example, when a highly competitive and less competitive species are mixed, more resources flow to the more competitive species. The yield gain of more competitive species from increases in resource availability (resources per capita) often exceeds the yield loss of less competitive species from decreases in resource availability (Huston, 1997; Montazeaud et al., 2018; Pillai and Gouhier, 2019; Tilman et al., 1997; Wardle, 1999), resulting in positive deviation from the total expected yield of both species (i.e., positive biodiversity effect). Competitive interactions can also result in negative biodiversity effects if the less competitive species suffers more than the more competitive species benefits.

The positive contribution of competitive dominance to ecosystem productivity based on the null expectation can be illustrated with species mixture data. In the situations of both simulated and experimental mixtures with two species (Huang et al., 2009; Mahaut et al. 2020), the partial density monoculture yield of the more competitive species exceeded the total expected yield of both species in 32 of the 35 mixtures (Table 1), meaning that a higher productivity expected for a given species mixture under the null hypothesis can be generally achieved by removing less competitive species from the mixture. In other words, the positive biodiversity effects of mixtures detected through the null expectation can be obtained from the monocultures of more competitive species at partial (lower) densities. These positive biodiversity effects, however, result from competition-induced changes in resource availability, i.e., the yield gain of more competitive species from all resources in the mixture exceeds the total yield loss expected for the less competitive species, and are not related to effects of positive species interactions. Competitive interactions are the predominant type of interspecific relationships in plants and occur in all mixtures where constituent species are competitively different (Goldberg and Barton, 1992; Pillai and Gouhier, 2019). Reporting positive biodiversity effects, without indication of possible contributions from competitive interactions, can lead to unrealistic expectations for effects of positive interactions.

Given the limits of the additive partitioning in deciphering mechanisms of changes in ecosystem productivity (Forrester and Pretzsch, 2015; Loreau and Hector, 2001; Petchey, 2003), some alternative methods have been explored to identify underlying mechanisms. This includes the use of metrics such as relative yield totals (Hector, 1998; Roscher and Schumacher, 2016) and transgressive overyielding or superyielding (Drake, 2003; Forrester and Pretzsch, 2015; Špaèková and Lepš, 2001), the multiplicative partitioning or statistical analyses to separate effects of species interactions and species composition (Chen et al., 2020; Jaillard et al., 2018; Kirwan et al., 2009), the identification of relative species influences via measurement of competitive interactions (Brun et al., 2022; Mahaut et al., 2020; Montès et al., 2008), or influences of species diversity, composition, and density on net biodiversity effects (Jaillard et al., 2018; Mahaut et al., 2020; Polley et al., 2003; Roscher and Schumacher, 2016; Tatsumi and Loreau, 2023). However, although these methods may indicate the presence of strong positive or competitive interactions or connections of the additive components with niche and fitness differences (Carroll et al., 2011; Godoy et al., 2020; Turnbull et al., 2013), they are not capable of quantifying changes in ecosystem productivity by species interactions at species or community level (Loreau et al., 2012).

Here, we present the competitive partitioning model, a new framework for assessing effects of species interactions on ecosystem productivity based on null expectations and competitive growth responses, the proportional changes of individual size (biomass or volume) from full density monocultures to mixtures expected from species differences in growth and competitive ability (Figure 1). Current biodiversity experiments are generally established based on a replacement design with equal density across mixtures and monocultures (Sackville Hamilton, 1994) and do not provide data for estimating competitive growth responses. Thus, we applied this new methodological approach to simulated data generated with the GYPSY model (Huang et al., 2009) for trembling aspen (Populus tremuloides Michx.) and white spruce (Picea glauca (Moench) Voss) tree mixtures and greenhouse experimental data from grassland mixtures (Mahaut et al., 2020) (Table 1). We used partial density (species density in mixture) monocultures and the competitive exclusion principle to determine maximum competitive growth responses and full density monoculture yields to measure species ability to compete for resources and achieve maximum competitive growth responses in mixture. Our objectives are to (1) demonstrate the power of this new framework in distinguishing effects of species interactions responsible for changes in ecosystem productivity, (2) illustrate how the mechanisms of changes in ecosystem productivity differ between the assessments with competitive and additive partitioning models, and (3) assess the effects of competitive interactions on ecosystem productivity relative to those of other species interactions.

Earlier version: Under the competitive hypothesis, more competitive species gain size (biomass or volume) and less competitive species lose size in mixture relative to their full density monocultures. The magnitude of competitive growth responses (i.e., proportional changes in individual size) increases with relative competitive ability (proportional deviations of species competitive ability from community average) from a minimum of 0 (at community average or null expectation) to a maximum of Wp/Wf – 1 for more competitive species and -1 for less competitive species at competitive exclusion. The changes may not be linear, greater near community average and smaller with deviations of competitive ability (Gaudet and Keddy, 1988) or recourse availability (e.g., light, see Brüllhardt et al., 2020) away from community averages. Wf represents for individual size in full density monocultures and Wp represents for individual size in partial density monocultures.

Revised version: Under the competitive hypothesis, more competitive species gain size (biomass or volume) and less competitive species lose size in mixture relative to their full density monocultures. The magnitude of competitive growth responses (i.e., proportional changes in individual size) increases with relative competitive ability (proportional deviations of species competitive ability from community average) from a minimum of 0 (at community average or null expectation) to a maximum of Wip/Wif – 1 for more competitive species and -1 for less competitive species at competitive exclusion. The changes may not be linear, greater near community average and smaller with deviations of competitive ability (Gaudet and Keddy, 1988) or resource availability (e.g., light, see Brüllhardt et al., 2020) away from community averages. Wif represents for individual size in full density monocultures and Wip represents for individual size in partial density monocultures.

Competitive expectation

Under the null hypothesis, species in mixture are assumed to be competitively equivalent (i.e., equal interspecific and intraspecific interactions) and species growth (or any other function) stays the same in mixture as in monocultures (Loreau and Hector, 2001). Deviations of observed yields (biomass or volume) from total null expectation therefore produce net biodiversity effects or changes in ecosystem productivity (Barry et al., 2019; Loreau and Hector, 2001). Thus, the net biodiversity effect is the difference between total observed yields ∑YOi and total null expectation ∑YNEi, calculated from species composition (or relative yield) in mixture RYEi and full density monoculture yield Mif.

To separate effects of competitive interactions from those of other species interactions, we would need the hypothesis that constituent species share an identical niche but differ in growth and competitive ability (i.e., absence of positive/negative interactions). Competitive growth responses, i.e., proportional changes in individual size (biomass or volume) expected from full density monocultures to mixtures under this competitive hypothesis, can be determined from species differences in growth and competitive ability (Figure 1). The magnitudes of competitive growth responses depend on: (1) species resource needs in full density monocultures; and (2) changes in resource availability (resources per capita) from full density monocultures to full density mixtures that are related to species composition and relative ability to compete for resources in mixture.

To estimate competitive growth response, we first determine the maximum competitive growth response MGi that species can reach in a scenario of competitive exclusion (Murray, 1986) where a more competitive species completely dominates and grows like a partial density monoculture, while a less competitive species is competitively eliminated. Here, wip stands for individual size (biomass or volume) in partial density (species density in mixture) monoculture and wif is for individual size in full density monoculture; both parameters are calculated from species monoculture yields and initial densities in partial/full density monocultures. A more competitive species gains in resource availability and in individual size, while a less competitive species loses 100% resource availability and size in mixture relative to full density monoculture (Figure 1). This simplified competitive dominance applies to mixtures where constituent species differ considerably in competitive ability, especially when competition for light is intense (Grime, 1973; Lieffers et al., 1993). In most cases, however, competitive exclusion would not occur in biodiversity experiments that are controlled in density, composition, duration, and scale. More competitive species are negatively affected by interspecific competition relative to their partial density monocultures (Holland and DeAngelis, 2009), while less competitive species are competitively suppressed relative to their null expectations, but not eliminated (Aarssen, 1989; Connolly et al., 2001; Mahaut et al., 2020; Montazeaud et al., 2018). Other than changes in resource availability, the estimated maximum competitive growth response would include the effects of density-dependent pests, pathogens, or microclimates.

Next, we use relative competitive ability RCi to represent species ability to compete for resources and achieve their maximum competitive growth response in mixture (Aarssen, 1989). RCi is defined by the deviation of species competitive ability Hi from community average (excluding species i) relative to community maximum Hm. Competitive ability is species ability to exploit resources and can be assessed by growth attributes such as full density monoculture yields in biomass or volume (Aarssen, 1989; Aschehoug et al., 2016; Gaudet and Keddy, 1988; Grace, 1990; Montazeaud et al., 2018). Greater competitive advantages would lead to increasingly positive competitive growth responses and larger competitive disadvantages would lead to increasingly negative competitive growth responses in mixture relative to full density monocultures. Species competitive expectation YCEi is the yield based on null expectation YNEi and competitive growth response MGiRCi derived from species differences in growth and competitive ability. The term MGiRCi would be positive for more competitive species MGi > 0, negative for less competitive species MGi < 0, and zero when species are at community average or competitively equivalent RCi =0 (Figure 1). At individual species level, the deviation of observed yield YOi from the competitive expectation YCEi can be positive (representing relaxation of interspecific competition from competitive expectation and therefore dominance of positive interactions), neutral (representing similar intra- and interspecific interactions or offset of positive and negative interactions), or negative (representing dominance of negative interactions). The difference between competitive and null expectations is competitive yield change resulting from competitive interactions. Isolating facilitative effect from those of other positive interactions (i.e., resource partitioning) is also possible for individual species by comparing species’ observed yield in mixture to their partial density monoculture yield, adapting the principle of additive design for biodiversity–ecosystem productivity experiments (Sackville Hamilton, 1994). At the community level, the difference in total yield between competitive expectations ∑YCEi and null expectations ∑YNEi of all species can be attributed to the competitive effect, whereas the difference between observed yields ∑YOi and competitive expectations ∑YCEi can be attributed to the dominance of positive effects when the difference >0 or to the dominance of negative effects when the difference <0. The community positive effect can be further partitioned by mechanisms of positive interactions (resource partitioning and facilitation), and facilitative effect can be classified as mutualism (+/+), commensalism (+/0), or parasitic (+/–) based on species specific assessments.

Case studies

Simulated aspen and spruce mixtures

Mixed forests composed of trembling aspen (Populus tremuloides Michx.) and white spruce (Picea glauca (Moench) Voss) are widely distributed in North America (Man and Lieffers, 1999). The two species, occurring in pure and mixed stands of varying compositions in natural conditions, generally have different canopy structure, rooting depth, leaf phenology, and requirements for light, moisture, and nutrients so are considered different in niche (Man and Lieffers, 1999).

We used a growth and yield simulation system GYPSY, developed for mixed species forests in western Canada (Huang et al., 2009), to generate data required for determining competitive expectation, i.e., mixtures and monocultures at partial and full densities. We were interested in knowing how effects of species interactions may vary with stand age and composition. We chose four ages: 20, 40, 60, and 80 years, to represent young, middle-aged, pre-mature, and mature stands, and five relative yields from nearly pure aspen (90% aspen), aspen-dominated (70% aspen), equal mixed (50% aspen), spruce-dominated (30% aspen), to nearly pure spruce (10% aspen) mixtures at average densities on fully stocked medium productivity sites (Peterson and Peterson, 1992) (Table 1). As GYPSY is developed from extensive field observations (Huang et al., 2009), the simulated data represent average yields in natural monocultures and mixtures of various composition and ages that may result from direct or indirect interactions between the two species (Barry et al., 2019; Man and Lieffers, 1999). Species competitive ability was approximated with stand volume in full density monocultures at site index 17 m and 14 m (height), respectively, for dominant and codominant aspen and spruce at 50 years of total age on medium productivity sites (Alberta Forest Service, 1985; Huang et al., 1994).

The partitioning of net biodiversity effects with competitive partitioning model indicated that the relative contribution of competitive interactions to net biodiversity effects decreased with age (Figure 2a), whereas community positive effects occurred generally in pre-mature and mature stands and more so in equal aspen and spruce mixtures (Figure 2b). Evidence for a facilitative effect (i.e., species observed yield in mixture > partial density monoculture yield) was only found on spruce in nearly pure mature spruce mixture (Table 1). Community negative effects occurred in young and middled-age mixtures where aspen did poorer than expected under the competitive expectation (Table S1). Averaged across all ages and compositions, the net biodiversity effect was 14.3 m3 ha-1 or 14% of the average full density monoculture yield. On a relative basis, 57% of the value (8.2 m3 ha-1) was attributable to positive effects, 65% (9.3 m3 ha- 1) to competitive effects and -23% (-3.2 m3 ha-1) to negative effects, compared to 63% and 37% for CE and SE with additive partitioning (Figure 2c, d).

Partitioning net biodiversity effects (changes in stand volume in mixtures relative to full density monocultures) with competitive (a and b) and additive (c and d) partitioning models for 20- to 80-year-old mixed trembling aspen and white spruce with varying species compositions from nearly pure aspen (90% aspen) to nearly pure spruce (90% spruce) on average medium productivity sites in western Canada. The growth and yield data (Table 1) are generated with GYPSY (Huang et al., 2009) and calculations of biodiversity components with this figure are detailed in Table S1.

Experimental grassland mixtures

Mahaut et al. (2020) conducted an experimental study in pots with six grassland species in full and half density monocultures, as well as in every combination of two, three and six species in mixture. The studied species vary in productivity and functional traits. We used aboveground biomass data from full and half density monocultures, as well as all combinations of two species in full density mixtures (Table 1). Species competitive ability was approximated with species biomass in full density monocultures.

Competitive partitioning suggested a weak facilitative effect in two mixtures that were associated with Lotus corniculatus, one of the two legume species studied (Tables 1 and S2). Community negative effects occurred in two mixtures with Sanguisorba minor and between the two legumes, due to negative effects of Sanguisorba minor on other species and Lotus corniculatus on Trifolium repens (Table S2). Across all 15 two-species mixtures, the average of net biodiversity effects on aboveground biomass was 3.9 g/pot or 49% of the full density monoculture yield averaged across all species. On a relative basis, 60% (2.3 g) of the value was attributable to positive effects, 45% (1.7 g) to competitive effects, and -5% (-0.2 g) to negative effects, compared to 84% and 16% for CE and SE with additive partitioning (Table 2).

Comparison of additive and competitive partitioning

The competitive partitioning model provides a conceptual framework that helps establish the linkage between changes in ecosystem productivity and effects of species interactions. Under this framework, in both simulated tree and experimental grassland mixtures, net biodiversity effects were generally positive, resulting from both positive and competitive interactions. The positive competitive effects result from greater competitive yield responses of more competitive species (Huston, 1997; Montazeaud et al., 2018; Tilman et al., 1997; Wardle, 1999), i.e., the yield gain of more competitive species from increases in resource availability exceeds the yield loss of less competitive species from decrease in resource availability, due to different competitive growth responses and full density monoculture yields (Tables S1 and S2). The greater competitive yield response of more competitive trembling aspen in young and middle-aged mixtures resulted in net biodiversity effects that were predominantly from competitive interactions (Figure 2a, b; Table S1). A facilitative effect was rarely detected, only on spruce in nearly pure spruce mature mixture, possibly due to improved nutrient availability by aspen (Man and Lieffers, 1999; Peterson and Peterson, 1992) and in mixtures with legume Lotus corniculatus in grassland mixtures, likely due to improved nitrogen availability (Mahaut et al., 2020). The strong negative effect of mixing on aspen (i.e., observed yields < competitive expectations, see Table S1) in young and middled-age mixtures suggests that trembling aspen originating naturally in high density does not benefit from mixing with white spruce. This is in contrast with white spruce where the deviations of observed yields from competitive expectations were generally positive, likely due to its shade tolerance and positive effects of aspen on nutrient availability (Man and Lieffers, 1999; Peterson and Peterson, 1992).

The net biodiversity effects reveal overall community yield changes from what would be expected from full density monoculture yields (Loreau and Hector, 2001, 2019). However, the underlying mechanisms of species interactions for detected changes are uncertain, due to effects of competitive interactions that can be positive or negative. In both simulated mixed forests and experimental grassland mixtures, the partial density monoculture yield of the more competitive species exceeded the total expected yield of both species in nearly all mixtures, resulting in positive biodiversity effects when positive species interactions are not involved. Such positive biodiversity effects may help understand changes in ecosystem productivity detected through the null expectation (Carroll et al., 2011; Godoy et al., 2020) but do not support the general expectation that positive changes in ecosystem productivity are associated with positive species interactions (Barry et al., 2019; Drake, 2003; Hooper et al., 2005; Loreau, 2000) or crop mixing in agriculture and forestry where interests are positive effects (Forrester and Pretzsch, 2015; Man and Lieffers, 1999; Montazeaud et al., 2018). Thus, partitioning net biodiversity effects by effects of species interactions (i.e., positive, competitive, negative effects) allows for estimating relative importance of mechanisms of species interactions responsible for changes in ecosystem productivity. This enables ecologists to determine which communities are truly benefitting from diversity in the form of positive interactions and which result from interspecies differences in growth and competitive ability. In the aspen and spruce mixtures examined, community positive effects occurred largely in pre-mature and mature mixtures (>=60 years), even though net biodiversity effects and CE were positive at nearly all stages of stand development (Figure 2). In grassland mixtures, two mixtures (i.e., Dactylis glomerata - Sanguisorba minor and Lotus corniculatus - Trifolium repens) had positive net biodiversity effects that resulted entirely from competitive interactions (Tables 2 and S2). These positive net biodiversity effects may indicate evidence for positive changes in ecosystem productivity when these changes were actually driven by differential outcomes of interspecific competition.

We demonstrate with simulated and experimental species mixtures that competitive interaction is a major source of the net biodiversity effects and affects the additive components of both CE and SE. In the scenario of partial density monocultures emulating a simplified scenario of competitive exclusion where positive net biodiversity effects resulted entirely from competitive interactions, CE was generally negative in tree mixtures where relative yield changes were smaller with more competitive species (negative mean relative yield change) but positive in grassland mixtures where relative yield changes were greater with more competitive species (positive mean relative yield change) (Tables S1 and S2). In aspen-spruce mixtures, SE was either larger (age 80) or smaller (age 20, 40, and 60) than competitive effect. In the grassland communities, SE was similar to competitive effect in 3 of the 15 mixtures (Bromus erectus - Lotus corniculatus, Dactylis glomerata - Lotus corniculatus, and Lotus corniculatus - Sanguisorba minor) but substantially smaller in the remaining mixtures (Table 2). Even in the extreme cases where SE was nearly zero (Bromus erectus - Sanguisorba minor and Bromus erectus - Trifolium repens), competitive effect was still relatively strong, suggesting that a common assumption to attribute positive net biodiversity effects to positive species interactions when SE is close to zero (Feng et al., 2022; Loreau and Hector, 2001) is not justified. Competitive interactions amplify interspecies differences and therefore SE, while species interference on more competitive species (Bromus erectus - Sanguisorba minor) and greater positive species interactions on less competitive species (Bromus erectus - Trifolium repens) (Table S2) reduce SE. Our findings, along with some recent suggestions (Bourrat et al., 2023; Carroll et al., 2011; Godoy et al., 2020; Turnbull et al., 2013), do not support some common exercises that attribute positive/negative CEs to positive/negative interactions (Barry et al., 2019; Fargione et al., 2007; Feng et al., 2022; Hagan et al., 2021; Loreau and Hector, 2001; Petchey, 2003) and positive SEs to competitive interactions (Loreau, 2000; Loreau and Hector, 2001; Montazeaud et al., 2018).

Several assumptions are adopted with the development of competitive partitioning model. First, we assumed that more productive species are more competitive, which was mostly true with the species mixtures examined in terms of species yields in mixture relative to their full density monoculture yields (Table 1). There were exceptions with Dactylis:Plantago and Plantago:Trifolium in grassland mixtures where less productive species Plantago based on full density monoculture yields performed better in mixtures, which may suggest positive effects on Plantago and negative effects on the other species. However, species yields in partial versus full density monocultures indicate that these exceptions resulted from the stagnation of more productive Plantago in full (high) density monocultures, due to restricted rooting volumes and supplies of water and nutrients in pots (Chung et al., 2021). The release of growth stagnation in partial density monocultures and mixtures affects not only the mechanisms of changes in ecosystem productivity assessed with competitive and additive partitioning models but also the magnitudes of the changes in ecosystem productivity detected through the null expectation (Loreau and Hector, 2001). In this case, the use of partial density monoculture yields in assessment of species competitive ability are probably more appropriate, which would increase the values of competitive effects and reduce those of positive effects based on the competitive expectations. Partial density monocultures help detect competition-related yield changes, which would be otherwise unknown. Second, we used partial density monocultures to determine maximum competitive growth responses in a simplified competitive exclusion where more competitive species are not affected and grow like partial density monocultures, while less competitive species are competitively eliminated. We assumed a linear relationship between competitive growth responses and species relative competitive ability (Figure 1), which should be a good approximation (Gaudet and Keddy, 1988), but may not be the most appropriate, e.g., light availability and plant height (Brüllhardt et al., 2020). Linear models can overestimate or underestimate competitive expectations of all species and affect individual species assessments but would have limited impacts on community level assessments, due to different directions of competitive growth responses among species.

Estimation of maximum competitive growth responses in mixtures of more than two species or in low densities may require additional efforts. In mixtures where more than one species are above average competitive ability, maximum increase in resource availability would be smaller, (j = species more competitive than species i), instead of Maximum competitive growth increases determined in partial density monocultures may need to be adjusted accordingly. For less competitive species, complete elimination would not occur in mixtures of extremely low densities or small individual size where growth is largely independent of resource availability. Proportional growth reduction (, opposite of proportional growth increase in partial density monocultures ) or partial elimination (by more competitive species) may be more appropriate for estimating maximum competitive growth reductions.

The competitive partitioning model incorporates effects of competitive interactions into the conventional null expectation and assists: (1) understanding the mechanisms of species interactions driving positive biodiversity-productivity relationships (Cardinale et al., 2012; Isbell et al., 2017) by relative contributions of positive interactions from more species of diverse niche (Hooper et al., 2005; Tilman et al., 1997) and competitive interactions from greater yield responses of more competitive and more productive species to changes in resource availability in mixture (Aarssen, 1997; Huston, 1997; Pillai and Gouhier, 2019; Tilman et al., 1997), (2) examining the mechanisms of key species in influencing ecosystem productivity by roles of competitive interactions due to their higher productivity and competitive advantages (Loreau and Hector, 2001; Mahaut et al., 2020), (3) meaningful comparisons of changes in ecosystem productivity (i.e., relative yield totals) across different ecosystems based on competitive expectations, and (4) redefining changes in ecosystem productivity such that positive changes result from positive interactions and negative changes from negative interactions (Drake, 2003; Hooper et al., 2005; Loreau, 2000). Under this competitive partitioning framework, null expectations are either raised for more competitive species with increased resource availability or lowered for less competitive species with decreased resource availability in mixture relative to full density monocultures. The magnitudes of competitive growth responses are derived from species relative competitive ability in mixture and their maximum competitive growth responses determined from full and partial density monocultures under a simplified scenario of competitive exclusion. Current biodiversity experiments do not have partial density monocultures and therefore do not provide estimates of competitive growth responses that would occur in mixture due to interspecies differences in growth and competitive ability. However, density-size/yield relationship is one of the most extensively studied areas in ecology (Watkinson, 1980; Weiner and Freckleton, 2010). Previous research can help determine maximum competitive growth responses through density-size/yield relationships (Huang et al., 2009). This means that the competitive partitioning model can be used to investigate changes in ecosystem productivity by effects of species interactions in current biodiversity experiments established with replacement design (Fargione et al., 2007; Loreau and Hector, 2001).

Conclusions

Competitive interactions are the major source of biodiversity effects and affect the additive components of both CE and SE. The interpretations of CEs and SEs with specific mechanisms of species interactions lack mathematical or ecological justifications. For example, attributing positive biodiversity effects solely to effects of positive interactions, as commonly seen in the literature (Barry et al., 2019; Fargione et al., 2007; Feng et al., 2022; Hooper et al., 2005; Loreau and Hector, 2001), assumes that competitive interactions are productivity-neutral (SE close to zero) (Feng et al., 2022; Loreau and Hector, 2001). This assumption is true if ecosystem productivity is defined by total resources pool (Tilman et al., 2014) but generally not true if mixture productivity is defined by species monoculture yields as with the additive partitioning (Loreau and Hector, 2001), due to the predominance of competitive interactions in interspecific relationships (Goldberg and Barton, 1992) and positive effects of competitive interactions on ecosystem productivity based on the null expectation.

The competitive partitioning model is based on ecological theories of interspecific relationships, competitive ability, and competitive exclusion principle and experimental designs of both replacement and additive series. This new partitioning approach enables meaningful assessments of species interactions at both species and community levels and provides detailed insights into the mechanisms of species interactions that drive changes in ecosystem productivity. We believe that our framework, which is admittedly perfectible, is one promising avenue to determine effects of species interactions responsible for changes in ecosystem productivity, an approach that is long sought after in biodiversity–ecosystem productivity research (Loreau and Hector, 2012, 2019; Mahaut et al., 2020).

Added in the revised version: Supporting information